Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name:

```python
NAME = "" # put your name here
```

---

# Lab5: Manipulating Spotify API Data Using Python `requests`

## Objective
In this lab, you will practice using Python to interact with the Spotify Web API. By the end of this lab, you will be able to:
- Retrieve data from the Spotify API using HTTP requests.
- Process and handle JSON data.
- Display specific information about artists, albums, and tracks.

## Instructions
Complete all cells marked as `YOUR CODE HERE`. Each `YOUR CODE HERE` cell has instructions on what is expected.

---


### Setup:

1. **Spotify Developer Account:**
   - Go to the [Spotify Developer Dashboard](https://developer.spotify.com/dashboard/login) and log in with your Spotify account.
   - Create a new application in the dashboard to obtain your `client_id` and `client_secret`.
   - These credentials are necessary to authenticate with Spotify's API.
   - Keep these credentials safe as you will use them to generate your access token.

**Note**: Remember not to share your API keys with anyone.

2. **Storing Credentials Securely in a `.env` File:**
   - Create a `.env` file in the root directory of your project. In this file, add your `client_id` and `client_secret` as shown below:
     ```
     SPOTIFY_CLIENT_ID=your_client_id_here
     SPOTIFY_CLIENT_SECRET=your_client_secret_here
     ```
   - Do **not** hardcode your credentials in the notebook. The notebook will load the credentials from the `.env` file using the `python-dotenv` library.

3. **Installing Required Libraries:**
   - Install the required libraries by running the following commands in your terminal:
     ```bash
     pip install requests python-dotenv
     ```

## Task 1: Authentication (OAuth 2.0)

Spotify's API requires OAuth 2.0 for authentication. 
In this task, you will request an access token using your `client_id` and `client_secret` from `.env` file.

```python
# Import required libraries
import requests  # to make HTTP requests
import json      # to handle JSON data
import base64    # for encoding client credentials
import pandas as pd  # for data processing
from dotenv import load_dotenv  # to load environment variables
import os  # to access environment variables

# Load the environment variables from the .env file
load_dotenv()

# Fetch credentials from environment variables
CLIENT_ID = os.getenv('SPOTIFY_CLIENT_ID')  # Load from .env file
CLIENT_SECRET = os.getenv('SPOTIFY_CLIENT_SECRET')  # Load from .env file
TOKEN_URL = 'https://accounts.spotify.com/api/token'

```

```python
def get_access_token(client_id, client_secret):
    """
    Function to retrieve the Spotify API access token using client credentials.
    
    Parameters:
        client_id (str): Spotify client ID
        client_secret (str): Spotify client secret
        
    Returns:
        str: Bearer token for accessing Spotify API
    """
    url = "https://accounts.spotify.com/api/token"
    headers = {
        'Content-Type': 'application/x-www-form-urlencoded'
    }
    data = {
        'grant_type': 'client_credentials',
        'client_id': client_id,
        'client_secret': client_secret
    }
    
    response = requests.post(url, headers=headers, data=data)
    if response.status_code == 200:
        # Parse and return the access token
        token = response.json().get('access_token')
        return token
    else:
        raise Exception(f"Failed to obtain access token. Status code: {response.status_code}")

# Example usage (Replace with actual client_id and client_secret from the .env file)
# token = get_access_token('your_client_id_here', 'your_client_secret_here')
# print(token)

```

<!-- grade_id: Authentication, points: 1, grade: true, solution: false -->
```python
# TEST CASE: Authentication

access_token = get_access_token(CLIENT_ID, CLIENT_SECRET)
assert access_token is not None  # Ensure the token is not None

```

```python

```

# Task 2: Searching for an Artist
In this task, you will implement a search function to look for an artist based on a search query.



<!-- grade_id: cell-ad3686c0a1fa46d8, points: 0, grade: false, solution: true -->
```python
#VALIDATION_FIELD[func] search_artist

def search_artist(token, query):
    """
    Function to search for an artist on Spotify.
    
    Parameters:
        token (str): OAuth token for Spotify API
        query (str): Search query (artist name)
        
    Returns:
        tuple: Artist name (str), Artist ID (str)
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
search_artist(access_token, "The Weeknd")
```

Expected Output: 
<table>
  <tr>
    <td>Result:</td>
    <td>('The Weeknd', '1Xyo4u8uXC1ZmMpatF05PJ')</td> 
  </tr>


<!-- grade_id: Search-for-Artist, points: 1, grade: true, solution: false -->
```python
# TEST CASE: Search for Artist

artist_name, artist_id = search_artist(access_token, "The Weeknd")
assert artist_name == "The Weeknd"  # Ensure the correct artist is returned
assert artist_id == '1Xyo4u8uXC1ZmMpatF05PJ'  # Ensure the correct artist ID is returned


```

# Task 3: Fetching Related Artists
For this task, you will fetch related artists to a given artist.

<!-- grade_id: cell-250c5cc13852de83, points: 0, grade: false, solution: true -->
```python
#VALIDATION_FIELD[func] get_related_artists

def get_related_artists(token, artist_id):
    """
    Function to retrieve related artists of a given artist.
    
    Parameters:
        token (str): OAuth token for Spotify API
        artist_id (str): Spotify Artist ID
        
    Returns:
        list: List of related artist names (str)
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
get_related_artists(access_token, '1Xyo4u8uXC1ZmMpatF05PJ')
```

Expected Output: 
<table>
  <tr>
    <td>Result like:</td>
    <td>['Gesaffelstein',
 'Lily-Rose Depp',
 'Swae Lee',
 'Doja Cat',
 'SZA',
 'Miguel',
 'G-Eazy',
 'The Neighbourhood',
 'French Montana',
 'Rihanna',
 '070 Shake',
 'Beach Weather',
 'Labrinth',
 'Justine Skye',
 'Ariana Grande',
 'Yoko Gold',
 'Ty Dolla $ign',
 'The Kid LAROI',
 'Chase Atlantic',
 'JENNIE'])</td> 
  </tr>


<!-- grade_id: Related-Artists, points: 1, grade: true, solution: false -->
```python
# TEST CASE: Related Artists

related_artists = get_related_artists(access_token, '1Xyo4u8uXC1ZmMpatF05PJ')  # The Weeknd's artist ID
assert len(related_artists) > 0  # Ensure that related artists are returned


```

```python

```

# Task 4: Find and Analyze a Show (Podcast)

In this task, you will search for a podcast show, retrieve metadata about the show and its episodes, and perform some analysis on the data.

Instructions:
1. Search for a Podcast Show: Use Spotify’s API to search for a podcast by name.
2. Retrieve Episode Details: Once the show is found, retrieve details about its episodes, such as release dates, durations, and descriptions.
3. Analysis: Perform an analysis, such as finding the longest episode or counting the episodes released in the last year.


<!-- grade_id: cell-d5f435ab1cacf6f9, points: 0, grade: false, solution: true -->
```python
#VALIDATION_FIELD[func] search_show

def search_show(token, query, limit=5):
    """
    Function to search for a podcast show by name.
    
    Parameters:
        token (str): OAuth token for Spotify API
        query (str): The name of the podcast show to search for
        limit (int): The number of shows to return (default is 1)
        
    Returns:
        tuple: (show_name, show_id) of the first result
    """
    url = 'https://api.spotify.com/v1/search'
    
    headers = {
        'Authorization': f'Bearer {token}'
    }
    
    params = {
        'q': query,
        'type': 'show',
        'limit': limit,
        'market': 'US'  # Specify the market

    }
    
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
search_show(access_token, "TED Talks Daily")
```

Expected Output: 
<table>
  <tr>
    <td>Result:</td>
    <td>('TED Talks Daily', '1VXcH8QHkjRcTCEd88U3ti')</td> 
  </tr>


<!-- grade_id: Search-Podcast-Show, points: 1, grade: true, solution: false -->
```python
# TEST CASE: Search for a Podcast Show

# Use your access_token to test the function (replace with valid access_token)
show_name, show_id = search_show(access_token, "The Joe Rogan Experience")
assert show_name is not None
assert show_id is not None


```

# Task 5: Get Episodes of a Podcast Show

Once you have the show ID, retrieve the first 10 episodes of the podcast, including the name, duration, release date, and description for each episode.

<!-- grade_id: cell-d0ed53210bc6d16e, points: 0, grade: false, solution: true -->
```python
#VALIDATION_FIELD[func] get_show_episodes

def get_show_episodes(token, show_id, limit=10):
    """
    Function to retrieve the episodes of a podcast show.
    
    Parameters:
        token (str): OAuth token for Spotify API
        show_id (str): The Spotify ID of the podcast show
        limit (int): The number of episodes to return (default is 10)
        
    Returns:
        list: List of dictionaries containing episode information (name, duration, release date)
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: Get-Podcast-Episodes, points: 1, grade: true, solution: false -->
```python
# TEST CASE: Get Podcast Episodes

episodes = get_show_episodes(access_token, show_id, limit=10)
assert len(episodes) == 10
assert 'name' in episodes[0]
assert 'duration_ms' in episodes[0]
assert 'release_date' in episodes[0]


```

# Task 6: Find the Longest Episode
In this task, you will analyze the episodes and find the longest one based on its duration.

<!-- grade_id: cell-4e4a06f5775b5dac, points: 0, grade: false, solution: true -->
```python
#VALIDATION_FIELD[func] find_longest_episode

def find_longest_episode(episodes):
    """
    Function to find the longest episode from a list of episodes.
    
    Parameters:
        episodes (list): List of dictionaries containing episode information (name, duration, release date)
        
    Returns:
        dict: The dictionary containing the longest episode's information
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: Find-Longest-Episode, points: 1, grade: true, solution: false -->
```python
# TEST CASE: Find Longest Episode

longest_episode = find_longest_episode(episodes)
assert longest_episode['name'] is not None


```

```python

```

```python

```